---
title: 'Manipulating Time'
description: 'Chapter description goes here.'
---

## Example coding exercise

```yaml
type: NormalExercise
key: e8c1edbe67
lang: python
xp: 100
skills: 2
```

This is an example exercise.

`@instructions`


`@hint`


`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}

```

---

## Functions That Return Date And Time

```yaml
type: VideoExercise
key: 0046cfeec8
xp: 50
```

`@projector_key`
019c08546179ae6a555ea3851cb638c0
